<img src="logo.png" height="200" width="900"> 

#  Сбор данных: грязная работа вашими руками 

Пришло время самостоятельно написать парсер! Мы будем собирать данные [о ценах на книги.](http://books.toscrape.com)

<!-- grade_id: cell-a1d057ec06e84bdd, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you
    
# Подгрузите все необходимые для работы пакеты.

# Если ваш код будет ругаться, что нет пакета lxml, установите его 
# Для этого выполните в одной из ячеек команду !pip3 install lxml
from typing import List
import requests
from bs4 import BeautifulSoup

import pandas as pd
from tqdm.notebook import tqdm
# your code here

```

```python
!pip3 install lxml
```

Output:
```
Collecting lxml
  Downloading lxml-4.6.3-cp37-cp37m-manylinux2014_x86_64.whl (6.3 MB)
[K     |████████████████████████████████| 6.3 MB 21.1 MB/s eta 0:00:01
[?25hInstalling collected packages: lxml
Successfully installed lxml-4.6.3
[33mWARNING: You are using pip version 20.2.4; however, version 21.2.1 is available.
You should consider upgrading via the '/opt/conda/bin/python3 -m pip install --upgrade pip' command.[0m
```

Прогуляйтесь на сайт http://books.toscrape.com/ и изучите его структуру.  


# 1. Сбор ссылок на книги

Напишите функцию `get_soup`, которая по ссылке возвращает html-разметку страницы в формате `bs4` 

<!-- grade_id: cell-bb3dd3de95e8b98e, points: 0, grade: false, solution: true -->
```python
def get_page_soup(url_link):
    
    ### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
    # will the code be with you
    response = requests.get(url_link)
    soup = BeautifulSoup(response.content)
    return soup

# your code here

```

```python
main_url = 'http://books.toscrape.com/catalogue/'
page_number = 5

soup = get_page_soup(main_url + f'page-{page_number}.html')
soup.find_all('a')[0].text
```

Output:
```
'Books to Scrape'
```

```python
main_url + f'page-{page_number}.html'
```

Output:
```
'http://books.toscrape.com/catalogue/page-5.html'
```

Напишите функцию `get_books_links`, которая находит в html-разметке страницы ссылки на странички с отдельными книгами. 

```python
for i in soup.find_all('article', {'class': 'product_pod'}):
    print(i.a.get('href'))
```

Output:
```
princess-jellyfish-2-in-1-omnibus-vol-01-princess-jellyfish-2-in-1-omnibus-1_920/index.html
princess-between-worlds-wide-awake-princess-5_919/index.html
pop-gun-war-volume-1-gift_918/index.html
political-suicide-missteps-peccadilloes-bad-calls-backroom-hijinx-sordid-pasts-rotten-breaks-and-just-plain-dumb-mistakes-in-the-annals-of-american-politics_917/index.html
patience_916/index.html
outcast-vol-1-a-darkness-surrounds-him-outcast-1_915/index.html
orange-the-complete-collection-1-orange-the-complete-collection-1_914/index.html
online-marketing-for-busy-authors-a-step-by-step-guide_913/index.html
on-a-midnight-clear_912/index.html
obsidian-lux-1_911/index.html
my-paris-kitchen-recipes-and-stories_910/index.html
masks-and-shadows_909/index.html
mama-tried-traditional-italian-cooking-for-the-screwed-crude-vegan-and-tattooed_908/index.html
lumberjanes-vol-2-friendship-to-the-max-lumberjanes-5-8_907/index.html
lumberjanes-vol-1-beware-the-kitten-holy-lumberjanes-1-4_906/index.html
lumberj
```

<!-- grade_id: cell-565bf344d89fef43, points: 0, grade: false, solution: true -->
```python
def get_book_links(page_soup) -> List[str]:
    
    ### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
    # will the code be with you
    links = []
    for i in soup.find_all('article', {'class': 'product_pod'}):
        links.append(i.a.get('href'))
    return links

# your code here

```

С помощью цикла соберите в лист `book_links` первые 200 книг.

<!-- grade_id: cell-883211fab847e672, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

book_links = []

for num in range(1, 11):
    soup = get_page_soup(main_url + f'page-{num}.html')
    book_links.extend(get_book_links(soup))

# your code here
book_links
```

Output:
```
['a-light-in-the-attic_1000/index.html',
 'tipping-the-velvet_999/index.html',
 'soumission_998/index.html',
 'sharp-objects_997/index.html',
 'sapiens-a-brief-history-of-humankind_996/index.html',
 'the-requiem-red_995/index.html',
 'the-dirty-little-secrets-of-getting-your-dream-job_994/index.html',
 'the-coming-woman-a-novel-based-on-the-life-of-the-infamous-feminist-victoria-woodhull_993/index.html',
 'the-boys-in-the-boat-nine-americans-and-their-epic-quest-for-gold-at-the-1936-berlin-olympics_992/index.html',
 'the-black-maria_991/index.html',
 'starving-hearts-triangular-trade-trilogy-1_990/index.html',
 'shakespeares-sonnets_989/index.html',
 'set-me-free_988/index.html',
 'scott-pilgrims-precious-little-life-scott-pilgrim-1_987/index.html',
 'rip-it-up-and-start-again_986/index.html',
 'our-band-could-be-your-life-scenes-from-the-american-indie-underground-1981-1991_985/index.html',
 'olio_984/index.html',
 'mesaerion-the-best-science-fiction-stories-1800-1849_983/index.html',
```

```python
print(get_page_soup(main_url + 'page-{}.html'.format(page_number)).text[:300])
```

Output:
```
All products | Books to Scrape - Sandbox

















Books to Scrape We love being scraped!








Home

All products









                            
                                Books
```

<!-- grade_id: cell-fb6eac2f95e1147e, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
assert len(book_links) == 200
```

# 2. Сбор информации о книгах 

Напишите несколько небольших функций, которые собирают различные данные об одной книге, необходимые для ответов на вопросы ниже. Информацию о книге собирайте в виде словаря вида 

```
{ 'name': 'Преступление и наказание', 'rating': 1, 'description': 'ужасно депрессивная книга', ... }

```

```python
main_url + book_links[0]
```

Output:
```
'http://books.toscrape.com/catalogue/a-light-in-the-attic_1000/index.html'
```

<!-- grade_id: cell-a3b3b4c91addb471, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here
def get_info_from_book(book_soup):
    ratings = {
        'One': 1,
        'Two': 2,
        'Three': 3, 
        'Four': 4,
        'Five': 5
    }
    rating = ratings[list(map(lambda x: x.get('class'), book_soup.find_all('p')))[2][1]]
    
    
    stats = book_soup.find('table', {'class': 'table table-striped'}).find_all('td')
    try:
        price = float(stats[2].text[1:])
    except:
        price = None
        
    try:
        tax = float(stats[4].text[1:])
    except:
        tax = None
        
    if len(book_soup.find_all('div', {'class': 'sub-header', 'id': 'product_description'})) == 0:
        description = None
    else:
        description = book_soup.find_all('p')[3].text
    return {
        'rating': rating,
        'price': price,
        'tax': tax,
        'description': description
           }
```

Пройдите циклом по всем сыслкам из списка `book_links` и соберите данные о книгах в вектор `book_info`. 

```python
data = pd.DataFrame({})
for book in tqdm(book_links):
    res = requests.get(main_url + book)
    sp = BeautifulSoup(res.content)
    data = data.append(get_info_from_book(sp), ignore_index=True)
```

Output:
```
HBox(children=(FloatProgress(value=0.0, max=200.0), HTML(value='')))
```

<!-- grade_id: cell-240b7f956c433b45, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

book_info = []

for book in tqdm(book_links):
    res = requests.get(main_url + book)
    sp = BeautifulSoup(res.content)
    book_info.append(get_info_from_book(sp))

# your code here

```

Output:
```
HBox(children=(FloatProgress(value=0.0, max=200.0), HTML(value='')))
```

Превратим вектор из информации в полноценную таблицу с данными. 

```python
book_info_df = pd.DataFrame(book_info)
print(book_info_df.shape)
book_info_df.head() 
```

Output:
```
(200, 4)
   rating  price  tax                                        description
0       3  51.77  0.0  It's hard to imagine a world without A Light i...
1       1  53.74  0.0  "Erotic and absorbing...Written with starling ...
2       1  50.10  0.0  Dans une France assez proche de la nôtre, un h...
3       4  47.82  0.0  WICKED above her hipbone, GIRL across her hear...
4       5  54.23  0.0  From a renowned historian comes a groundbreaki...
```

Теперь, когда все данные собраны, настало время ответить на несколько вопросов:

- У скольких книг отсутствует описание? Положите число, получившееся в результате ваших манипуляций с таблицей, в переменную `ans1`. 

```python
book_info_df.description.isna().sum()
```

Output:
```
1
```

<!-- grade_id: cell-818d91b73f034e1a, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

ans1 = book_info_df.description.isna().sum() 

# your code here

```

<!-- grade_id: cell-4b9e7e39b8921ba5, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
```

- Сколько раз в данных встречается налог, больший нуля?  Положите число, получившееся в результате ваших манипуляций с таблицей, в переменную `ans2`. 

```python
(book_info_df.tax > 0).sum()
```

Output:
```
0
```

<!-- grade_id: cell-be342205fa946a79, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

ans2 = (book_info_df.tax > 0).sum() 

# your code here

```

<!-- grade_id: cell-c3b1e5331d25a061, points: 0, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
```

- Сколько раз рейтинг книги составлял пять звезд? Положите число, получившееся в результате ваших манипуляций с таблицей, в переменную `ans3`. 

```python
(book_info_df.rating == 5).sum()
```

Output:
```
37
```

<!-- grade_id: cell-a904ff9c737f3bae, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

ans3 = (book_info_df.rating == 5).sum() 

# your code here

```

<!-- grade_id: cell-ac096895d662ffc8, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно

```

- Какова средняя цена книг (без учета налога)? Положите число, получившееся в результате ваших манипуляций с таблицей, в переменную `ans4`. 

```python
(book_info_df.price).mean()
```

Output:
```
34.79624999999999
```

<!-- grade_id: cell-60c854eb83c66df8, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

ans4 = (book_info_df.price).mean() 

# your code here

```

<!-- grade_id: cell-36295b0dc83ec0ad, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно

```